# Applied Machine Learning (2022), exercises


## General instructions for all exercises

Follow the instructions and fill in your solution under the line marked by tag

> YOUR CODE HERE

Remove also line 

> raise NotImplementedError()

**Do not change other areas of the document**, since it may disturb the autograding of your results!
  
Having written the answer, execute the code cell by and pressing `Shift-Enter` key combination. The code is run, and it may print some information under the code cell. The focus automatically moves to the next cell and you may "execute" that cell by pressing `Shift-Enter` again, until you have reached the code cell which tests your solution. Execute that and follow the feedback. Usually it either says that the solution seems acceptable, or reports some errors. You can go back to your solution, modify it and repeat everything until you are satisfied. Then proceed to the next task.
   
Repeat the process for all tasks.

The notebook may also contain manually graded answers. Write your manually graded answer under the line marked by tag:

> YOUR ANSWER HERE

Manually graded tasks are text in markdown format. It may contain text, pseudocode, or mathematical formulas. You can write formulas with $\LaTeX$-syntax by enclosing the formula with dollar signs (`$`), for example `$f(x)=2 \pi / \alpha$`, will produce $f(x)=2 \pi / \alpha$

When you have passed the tests in the notebook, and you are ready to submit your solutions, download the whole notebook, using menu `File -> Download as -> Notebook (.ipynb)`. Save the file in your hard disk, and submit it in [Moodle](https://moodle.uwasa.fi) or EUNICE Moodle under the corresponding excercise.

Your solution should be an executable Python code. Use the code already existing as an example of Python programing and read more from the numerous Python programming material from the Internet if necessary. 


```python

```

## Wine quality determination (Regression)
A chemical analysis was carried out for 1599 red wine samples, after which the quality of each sample was analyzed by experts in scale 0..10.

The chemical analaysis reveals 11 features for each wine sample, which are:
fixed acidity, volatile acidity, citric acid, residual sugar, chlorides, free sulfur dioxide, total sulfur dioxide, density, pH, sulphates, and alcohol (concentration). Each feature is a floating point value, and the quality is an integer between 0 to 10.  


## Task 1


### Read the data

- Read a CSV-file, called `winequality-red.csv`
- Separate 11 first columns and all rows to your design matrix X
- Use the last column, quality, as dependent variable (target)


### Make training set and test set
Separate your data X and y to training set (X_train, y_train) which contains 75% of the data and to the test set X_test, y_test which contains 25% of the data.

```python
##>>> Some code for bootstrap
import matplotlib.pyplot as plt
import pandas as pd
import numpy as np
```

<!-- grade_id: cell-7ccfa9a73a1d6f54, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-813c7fc121f0d455, points: 2, grade: true, solution: false -->
```python
points=0

## >> Some tests, do not change
if (X_train.shape==(1199,11)):
    points+=1
if (X_test.shape==(400,11)):
    points+=1

points
```

## Task 2

- Study the data, select a regression algorithm for predicting the quality of the wine, based on it's chemical features. You can assume that the quality is a floating point number.
- Train the regression algorithm using the training data.
- Use cross validation to test the performance of the regressor and tune it's parameters as good as you can
- Finally test the regressor with the test set
- Report the score (= $R^2$ = coefficient of determination) of the regressor in the training set, cross validation and in the test set
- Plot the predicted quality against the known quality 
- What does the $R^2$ score tells? 
- What is your opinion of the performance? Is there signs of overfitting?

<!-- grade_id: cell-93a0d47f20379f85, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import Pipeline

# YOUR CODE HERE
raise NotImplementedError()

print("Accurary in the training set..%f" % train_score)
print("Accurary in crossvalidation...%f" % cv_score)
print("Accurary in the test set......%f" % test_score)

```

<!-- grade_id: cell-04912d6dd5d188c3, points: 4, grade: true, solution: false -->
```python
### >>>> Some testing
# The coefficient of determination should be between 0.4 - 0.5 in the
# test set depending on your regression algorithm and parameters
points=0
assert('train_score' in globals()), "Please use train_score variable for storing the training score"
assert('cv_score' in globals()), "Please use cv_score variable for storing the crossvalidation score"
assert('test_score' in globals()), "Please use test_score variable for storing the test score"
points=1

assert(train_score>0.85), "Try again, if you can get better accuracy by tuning the regressor parameters"
points+=1


points
```

## Task 2 b

1. What does the $R^2$ score tells? 
2. What is your opinion of the performance? Is there signs of overfitting?

<!-- grade_id: cell-f35665f6a7dce240, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Task 3, Select the best features

Like it often is, some features are more important for regression than the others. 

- Study which features are the most important for predicting the quality. If you used Extratrees regressor or GradientBoostedTree regressor, then you can use the .feature_importances vector. Another option is to use LASSO or Elastic net models and SelectFromModel to reduce unimportant features.
- Plot a bar graph or stem plot to visualize the importance of features
- To check the importance of a single variable, plot a scatter plot using true quality in x-axis and the value of teh most important variable in y-axis.
- Make a vector `features` which contains the names of the three most important features in decreasing order

<!-- grade_id: cell-515bf1cb9ed3beee, points: 0, grade: false, solution: true -->
```python
features=['f1', ]
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-35fdf7b370e0abb8, points: 2, grade: true, solution: false -->
```python
assert(len(features)==3), "Please put the names of three most important features in list features!"
points=0.5

points
```

```python

```