# Unsupervised learning, clustering

## Task 1: Apply k-means

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
sns.set()  # for plot styling
```

The data-file `liver-spectroscopy.tab` is a text file which contains a Near Infrared Spectral (NIR) measurements of different cells from a liver. The file includes 731 samples, and each one has 234 variables. The variables describes the amount of infrared radiation absorbed in each of the tested 234 wavelengths. The last column in the data is the cell type, which can be one of the four types: 'collagen', 'glycogen', 'lipids' or 'DNA'. The column names are the wavelengths used in measuring the absorbtion values for the specific columns.

```python
D = pd.read_table('liver-spectroscopy.tab')

# Print the classes
print(D.type.unique())

# Separate the true classes out from the data, into separate variable
types=pd.Categorical(D.type)
del(D['type'])
```

Output:
```
['collagen' 'glycogen' 'lipids' 'DNA']
```

First task is to apply PCA to transform the data `D` to PCA projection `projected`. Then apply KMeans to `projected` and try to find clusters from the data. Use just enough PCA-components to contain more than 90% of the variance. Name your KMeans object as `kmeans`.

Plot the clusters using PC1 and PC2 axis using scatter plot. Use different color for each cluster in the plot.

<!-- grade_id: cell-10342f7d59a7a427, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[1], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-d2b7ace3c3ddd431, points: 2, grade: true, solution: false -->
```python
points=0
### BEGIN HIDDEN TESTS 
if (kmeans.cluster_centers_.shape==(4,3)) and (projected.shape==(731,3)):
    points+=1
### END HIDDEN TESTS 
if (abs(kmeans.inertia_-75)<5):
    points+=1
points
```

Output:
```
2
```

## Task 2, apply PCA and GMM

Use previous PCA projection and apply Gaussian Mixture Model to projected data and try to find clusters from the data. Use `gmm` as a name for GaussianMixture object.

Plot the clusters with different colors in the scatter plot using PC1 and PC2 axis.

<!-- grade_id: cell-3083b4a641efcef3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[2], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-116468ea10bd494e, points: 2, grade: true, solution: false -->
```python
points=0
### BEGIN HIDDEN TESTS 
if (projected.shape==(731,3)):
    points+=1
if (abs(gmm.bic(projected)+269)<10):
    points+=1
### END HIDDEN TESTS 
points
```

Output:
```
2
```

## Task 3: Confusion matrix

Calculate the confusion matrix `CM` between the clusters and true tissue types. Find from `CM`, the number of the cluster where *DNA* tissue type samples are most often assigned. Assign in variable `nDNA` the number of times the samples of tissue type *DNA* were assigned into this cluster.

Some instructions

1. Find the category number of DNA ny listing `types.categories`. The category number is the index of string 'DNA' in the list of categories.
1. Use `CM=confusion_matrix)=` -function from `sklearn.metrics` to plot the confusion matrix
1. Assign the found number in variable `nDNA`


```python
types.categories
```

Output:
```
Index(['DNA', 'collagen', 'glycogen', 'lipids'], dtype='object')
```

<!-- grade_id: cell-d003e438a21b9633, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[3], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-bb3792e534536dd9, points: 2, grade: true, solution: false -->
```python
points=0
assert('CM' in globals()), "Define the confusion matrix as CM please!"
assert(type(CM) == np.ndarray), "Confusion matrix is not an numpy array??"
assert('nDNA' in globals()), "Assign nDNA as istructed, please"

### BEGIN HIDDEN TESTS 
if(abs(nDNA-91)<5):
    points+=1
if (abs(CM.max(axis=0).sum()-706)<50):
    points+=1
### END HIDDEN TESTS 
points
```

Output:
```
2
```

## Task 3: The probability of the samples

The GMM model includes the function called `.predict_proba()` which returns the probability that a certain sample belongs to different clusters. Calculate the probabilities of each sample belongin to each cluster and assign the result in variable `P`. Then find out the proability that the last sample belongs in the cluster where most DNA samples belong to, and assign that probability in variable `pDNA`.

<!-- grade_id: cell-73bd925190782d71, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[4], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-ac22c0d94b014cc8, points: 2, grade: true, solution: false -->
```python
points=0
assert('P' in globals()), "Define probabilities, P as instructed, please."
assert('pDNA' in globals()), "Define probability, pDNA as instructed, please."
points+=1
### BEGIN HIDDEN TESTS 
assert(abs(pDNA-0.995)<0.001)
points+=1
### END HIDDEN TESTS 
points
```

Output:
```
2
```

## Task 4: Evaluation

1. Which clustering method, KMeans or GMM is better for this case and why?
1. How can the quality of the clustering methods be assessed if 
    1. the true classes are known?
    1. the true classes are not known?
    
Answer by writing text or Markdown text in the cell below.

<!-- grade_id: cell-e99622156ca1efab, points: 2, grade: true, solution: true -->
# YOUR ANSWER HERE