# Applied Machine Learning (2022), exercises


## General instructions for all exercises

Follow the instructions and fill in your solution under the line marked by tag

> YOUR CODE HERE

Remove also line 

> raise NotImplementedError()

**Do not change other areas of the document**, since it may disturb the autograding of your results!
  
Having written the answer, execute the code cell by and pressing `Shift-Enter` key combination. The code is run, and it may print some information under the code cell. The focus automatically moves to the next cell and you may "execute" that cell by pressing `Shift-Enter` again, until you have reached the code cell which tests your solution. Execute that and follow the feedback. Usually it either says that the solution seems acceptable, or reports some errors. You can go back to your solution, modify it and repeat everything until you are satisfied. Then proceed to the next task.
   
Repeat the process for all tasks.

The notebook may also contain manually graded answers. Write your manually graded answer under the line marked by tag:

> YOUR ANSWER HERE

Manually graded tasks are text in markdown format. It may contain text, pseudocode, or mathematical formulas. You can write formulas with $\LaTeX$-syntax by enclosing the formula with dollar signs (`$`), for example `$f(x)=2 \pi / \alpha$`, will produce $f(x)=2 \pi / \alpha$

When you have passed the tests in the notebook, and you are ready to submit your solutions, download the whole notebook, using menu `File -> Download as -> Notebook (.ipynb)`. Save the file in your hard disk, and submit it in [Moodle](https://moodle.uwasa.fi) or EUNICE Moodle under the corresponding excercise.

Your solution should be an executable Python code. Use the code already existing as an example of Python programing and read more from the numerous Python programming material from the Internet if necessary. 


```python

```

## Neural Networks

## This exercise is optional
If you decide to do this, write in the following cell, which of the exercises 1-9 you want to replace with this one.

<!-- grade_id: cell-89fbf057987d9e1e, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

The following code creates a data set is a 3D data set of values in design matrix X and corresponding ground truth in vector y.

```python
##>>> Some code for bootstrap
import matplotlib.pyplot as plt
import pandas as pd
import numpy as np

from sklearn import datasets 
from mpl_toolkits.mplot3d import Axes3D
```

```python
X,y=datasets.make_s_curve(n_samples=500, noise=0.2, random_state=0)

fig = plt.figure(figsize=(7,7))
ax = fig.add_subplot(111, projection='3d')
ax.scatter(X[:,0], X[:,1], X[:,2], c=y, cmap='rainbow')
ax.set_xlabel('x'); ax.set_ylabel('y'); ax.set_zlabel('z')
plt.title('Data in 3D space')

fig,axs=plt.subplots(nrows=1, ncols=3, figsize=(15,5))
axs[0].scatter(X[:,0], X[:,1], c=y, cmap='rainbow')
axs[1].scatter(X[:,0], X[:,2], c=y, cmap='rainbow')
axs[2].scatter(X[:,1], X[:,2], c=y, cmap='rainbow')

for i in range(3):
    axs[i].set_title('2D-projection')
    axs[i].set_xlabel("XXY"[i])
    axs[i].set_ylabel("YZZ"[i])
```

### Task 1: MLPRegressor

Implement a `MLPRegressor` to predict `y` using data in `X`. Try different number of hidden layers and number of perceptrons in the hidden layer to reach good prediction accuracy measured by `r2_score` and `cross_val_score`. Avoid too high complexity so that prediction in the training set would be significantly better than in the cross-validation set.   

```python
from sklearn.neural_network import MLPRegressor
from sklearn.metrics import r2_score
from sklearn.model_selection import cross_val_score
```

<!-- grade_id: cell-0fc6f13698a9a11b, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

print("Prediction accuracy in training set is", accuracy)
print("Prediction accuracy in cv is          ", accuracy_cv)
```

<!-- grade_id: cell-d8c91c52d6353031, points: 3, grade: true, solution: false -->
```python
if 'predictor' not in globals():
    print("Use name predictor for MLPRegressor object, please.")
assert(accuracy>0.95)
assert(accuracy_cv>0.93)
assert(type(predictor)==MLPRegressor)
```

## Task 2: Compare with Other regressor

Select some other suitable regressor, and compare it's accuracy with `MPLRegerssor`. Try to avoid overfitting.

<!-- grade_id: cell-7d99a0aa4cb932e4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9bb23c7dc18db6ff, points: 3, grade: true, solution: false -->
```python
if 'predictor2' not in globals():
    print("Use name predictor2 for Extratreesegressor object, please.")
assert(type(predictor2)==ExtraTreesRegressor)

```

## Task 3

- Compare the performance of two regressors
  - Accuracies
  - Training time
  - Prediction time (you may speculate this theoretically, or try to measure)
- Do you think Ridge regression, Lasso or Elastic net would work in this case? Why/WHy not?

<!-- grade_id: cell-93a0d47f20379f85, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE