## Soil moisture prediction (Regression)
A remote sensing system provides spectral images of soil, which is used to predict the soil moisture. The spectral data contains a matrix of reflected ligth at different wavelengts from one location. At the same time, the true soil moisture and temperature were measured with an accurate reference method. 

## Task 1


### Read the data

- Read a CSV-file, called `soilmoisture_dataset.csv`
- Extract all the columns which has a wavelength number as column header to your design matrix X
- Use the `soil_moisture` -column as dependent variable, y


### Make training set and test set
Separate your data X and y to training set (X_train, y_train) which contains 75% of the data and to the test set X_test, y_test which contains 25% of the data.

```python
import matplotlib.pyplot as plt
import pandas as pd
import numpy as np
```

<!-- grade_id: cell-7ccfa9a73a1d6f54, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[2], line 2
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-813c7fc121f0d455, points: 2, grade: true, solution: false -->
```python
errorfound=False
require=('X', 'y')
for var in require:
    if var not in globals():
        print(var, "not found! Please define it.")
        errorfound=True

### BEGIN HIDDEN TESTS
print(X_train.shape, X_test.shape, y_train.shape, y_test.shape, D.shape)
display(X.head())

if errorfound:
    points=0
else:
    points=1
if (X_train.shape[1]==125) and (abs(X_train.shape[0]-509)<10):
    points+=0.5
if (X_test.shape[1]==125):
    points+=0.5
points
### END HIDDEN TESTS
```

Output:
```
(509, 125) (170, 125) (509,) (170,) (679, 128)
                          454       458       462       466       470  \
datetime                                                                
2017-05-23 14:06:17  0.082131  0.055863  0.050047  0.047925  0.047498   
2017-05-23 14:08:17  0.079510  0.055326  0.049116  0.047579  0.046745   
2017-05-23 14:10:17  0.080599  0.054065  0.049150  0.047537  0.046525   
2017-05-23 14:12:17  0.078024  0.054972  0.049052  0.047920  0.046883   
2017-05-23 14:14:17  0.079973  0.055335  0.049314  0.047373  0.047028   

                          474       478       482       486       490  ...  \
datetime                                                               ...   
2017-05-23 14:06:17  0.046464  0.046669  0.046834  0.047469  0.048608  ...   
2017-05-23 14:08:17  0.046833  0.046329  0.047021  0.047660  0.048272  ...   
2017-05-23 14:10:17  0.046032  0.046276  0.046886  0.047230  0.048611  ...   
2017-05-23 14:12:17  0.046775  0.046835  0.046996  
```

## Task 2

- Study the data, select a regression algorithm for predicting the moisture
- Train the regression algorithm using the training data.
- Use cross validation to test the performance of the regressor and tune it's parameters
- Finally test the regressor with the test set
- Report the score (= $R^2$ = coefficient of determination) of the regressor in the training set, cross validation and in the test set
- Plot the predicted moisture against the known moisture
- What does the $R^2$ score tells? 
- What is your opinion of the performance? Is there signs of overfitting?

<!-- grade_id: cell-93a0d47f20379f85, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import Pipeline

# YOUR CODE HERE
raise NotImplementedError()


print("Accurary in the training set..%f" % train_score)
print("Accurary in crossvalidation...%f" % cv_score)
print("Accurary in the test set......%f" % test_score)

```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[3], line 4
      1 from sklearn.pipeline import Pipeline
      3 # YOUR CODE HERE
----> 4 raise NotImplementedError()
      7 print("Accurary in the training set..%f" % train_score)
      8 print("Accurary in crossvalidation...%f" % cv_score)

NotImplementedError:
```

<!-- grade_id: cell-04912d6dd5d188c3, points: 4, grade: true, solution: false -->
```python
# The coefficient of determination should be above 0.9 in the
# test set depending on your regression algorithm and parameters

errorfound=False
require=('train_score', 'cv_score', 'test_score')
for var in require:
    if var not in globals():
        print(var, "not found! Please define it.")
        errorfound=True

assert(train_score>0.85), "Try again, if you can get better accuracy by tuning the regressor parameters"

### BEGIN HIDDEN TESTS
if errorfound:
    points=0
else:
    points=2
if (test_score>0.9):
    points+=1
if (cv_score>0.9):
    points+=1    
### END HIDDEN TESTS

points
```

Output:
```
4
```

## Task 3, Select the best features

Like it often is, some features are more important for regression than the others. 

- Study which features are the most important for predicting the moisture. If you used Extratrees regressor or GradientBoostedTree regressor, then you can use the .feature_importances vector. Another option is to use LASSO or Elastic net models and SelectFromModel to reduce unimportant features.
- Plot a bar graph to visualize the importance of features
- To check the importance of a single variable, plot a scatter plot using true moisture in x-axis and the value of the most important variable in y-axis.
- Study the shape of the dependency. Is it linear or nonlinear?
- Make a vector `features` which contains the names of the 10 most important features in decreasing order

Name your variables according to the following table, in order to make the autograder tests to work:

| Variable | Name |
| -------- | ------|
| Best features | `features` |
| Barplot | `barplot` |
| Scatter plot  | `scatterplot` |

<!-- grade_id: cell-515bf1cb9ed3beee, points: 0, grade: false, solution: true -->
```python
features=[]

# You may find these lines usefull, in making the barplot
#x=range(len(i))
#barplot=plt.bar(....)
#plt.xticks(x, X.columns[i], rotation='vertical');

# YOUR CODE HERE
raise NotImplementedError()

```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Cell In[4], line 9
      1 features=[]
      3 # You may find these lines usefull, in making the barplot
      4 #x=range(len(i))
      5 #barplot=plt.bar(....)
      6 #plt.xticks(x, X.columns[i], rotation='vertical');
      7 
      8 # YOUR CODE HERE
----> 9 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-35fdf7b370e0abb8, points: 4, grade: true, solution: false -->
```python
errorfound=False
require=('features', 'barplot', 'scatterplot')
for var in require:
    if var not in globals():
        print(var, "not found! Please define it.")
        errorfound=True

assert(len(features)==10), "Please put the names of three most important features in list features!"

### BEGIN HIDDEN TESTS
if errorfound:
    points=0
else:
    points=1
    
print(features)

if scatterplot.draw:
    points+=1
if barplot.index:
    points+=1

for feature in features:
    if feature.startswith('8'):
        print("Yes", feature)
        points+=1
        break
### END HIDDEN TESTS
points
```

Output:
```
barplot not found! Please define it.
scatterplot not found! Please define it.
AssertionError                            Traceback (most recent call last)
Cell In[5], line 8
      5         print(var, "not found! Please define it.")
      6         errorfound=True
----> 8 assert(len(features)==10), "Please put the names of three most important features in list features!"
     10 ### BEGIN HIDDEN TESTS
     11 if errorfound:

AssertionError: Please put the names of three most important features in list features!
```

```python

```