# Applied Machine Learning (2023), exercises


## General instructions for all exercises

Follow the instructions and fill in your solution under the line marked by tag

> YOUR CODE HERE

Remove also line 

> raise NotImplementedError()

**Do not change other areas of the document**, since it may disturb the autograding of your results!
  
Having written the answer, execute the code cell by and pressing `Shift-Enter` key combination. The code is run, and it may print some information under the code cell. The focus automatically moves to the next cell and you may "execute" that cell by pressing `Shift-Enter` again, until you have reached the code cell which tests your solution. Execute that and follow the feedback. Usually it either says that the solution seems acceptable, or reports some errors. You can go back to your solution, modify it and repeat everything until you are satisfied. Then proceed to the next task.
   
Repeat the process for all tasks.

The notebook may also contain manually graded answers. Write your manually graded answer under the line marked by tag:

> YOUR ANSWER HERE

Manually graded tasks are text in markdown format. It may contain text, pseudocode, or mathematical formulas. You can write formulas with $\LaTeX$-syntax by enclosing the formula with dollar signs (`$`), for example `$f(x)=2 \pi / \alpha$`, will produce $f(x)=2 \pi / \alpha$

When you have passed the tests in the notebook, and you are ready to submit your solutions, validate and submit your solution using the nbgrader tools from the `Nbgrader/Assignment List`-menu.


```python

```

# Unsupervised learning, projections of multivariate data

## Task 1: Read and visualize spectral data

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
sns.set()  # for plot styling
```

The data-file `liver-spectroscopy.tab` is a text file which contains a Near Infrared Spectral (NIR) measurements of different cells from a liver. The file includes 731 samples, and each one has 234 variables. The variables describes the amount of infrared radiation absorbed in each of the tested 234 wavelengths. The last column in the data is the cell type, which can be one of the four types: 'collagen', 'glycogen', 'lipids' or 'DNA'. The column names are the wavelengths used in measuring the absorbtion values for the specific columns.

The data file is in text format, in which the columns are separated by tab-characters. 

The first task is to read the data file into a pandas dataframe called `D`. Use pandas function `.read_table()` for reading it.

<!-- grade_id: cell-40890b30971c4753, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c3377548c6d90983, points: 1, grade: true, solution: false -->
```python
assert('D' in globals()), \
    print("Use name D for NIR data, please.")    
assert(type(D)==pd.DataFrame)

assert(D.shape==(731,235)), "The shape of the data matrix is not quite right. Try to read it again"

```

Then visualize the spectral data by carrying out the following steps
1. Parse the `type` column as categorical value vector called `types` using function `pd.Categorical()`
1. Delete the original column 'type' from the dataframe (`del D['type']`)
1. parse the column names (`D.columns`) into a numerical vector of wavelengts, and name it as `wl`. You may find the `map()` function usefull for this purpose. Check from the last exercise how it was used. Round the wavelengths to integers, if you find a way to do that, but you can also used it as floating point values.
1. Assing the parsed `wl` vector as new column names of D (you can just assign it to `D.columns`)
1. Plot the first 10 samples. Select the first ten rows using `D.iloc[]`, and plot the transpose of them, using a transpose operator `.T`, as follows `D.iloc[].T.plot()`. Transpose is needed, because the plot function always places the row index in the x-axis. If the dataframe is first transposed, the column index (the wavelengths) becomes the row index, and plot will be correct.

You can also use `plt.xlabel()` and `plt.ylabel()` -functions to define labels for x- and y-axes.

The result should look like this:
![spectra](spectra.png)


<!-- grade_id: cell-912ad21a1457e9d0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d58036ed979386ec, points: 2, grade: true, solution: false -->
```python
assert(type(types)==pd.Categorical)
assert(len(types)==731)
assert(len(wl)==234)
assert(D.shape==(731,234))

```

## Task 2, calculate PCA projection and visualize

Apply PCA to the dataframe `D` and plot how much variance is explained by the first 10 principal components. Store the transformed `D` as variable `projectedD`.

Store the explained variance ratio in variable `varexp`

<!-- grade_id: cell-c9dd35cb35f900b2, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1f529977eb1b3300, points: 2, grade: true, solution: false -->
```python
assert('varexp' in globals()), \
    print("Use the name varexp for storing the explained variance ratio, please.")
assert(projectedD.shape==(731,10))

```

How many PCA components you would select if you need to explain at least 90% of the variance of the data?
Answer by assigning correct value to n below.

<!-- grade_id: cell-207be29faf366be6, points: 0, grade: false, solution: true -->
```python
n=0
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bd50285e72f560c6, points: 1, grade: true, solution: false -->
```python
assert('n' in globals()), \
    print("Use the name n for the number of principal components needed, please.")

```

### Visualize the data in PCA domain

Plot a scatter plot of two first principal components. Since `projectedD` is a numpy array and not a DataFrame, use the plotting function  `plt.scatter()`. The function requires two parameters, x-data and y-data, but some others are also usefull:
1. The first parameter is the x-data, give it the first principal compoent
1. The second parameter is the y-data, give it the second principal component
1. Give the class category for point colors `c=types.codes` 
1. Use the colormap parameter to use better colormap, e.g. `cmap='rainbow'`
1. Use rather small point size e.g. `s=10`
1. Make points somewhat transparent `alpha=0.5`

You can also use `plt.xlabel()` and `plt.ylabel()` functions to provide names for x and y axes, and `plt.colorbar()` to plot the colorbar explaining the meaning of the colors.
The plot should look like the following:

![pca](pca.png)

Notice that the different tissue types are clearly different already in two dimensional space.

<!-- grade_id: cell-26b0dd702df55434, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9216a6fb7de8c178, points: 2, grade: true, solution: false -->
```python

```

## Task 3, calculate UMAP manifold projection and visualize

Repeat the previously made projection using UMAP and visualize it. Save your UMAP projected data as `umapD`. You can use just two components for UMAP.

The result could look like the following t-SNE projection plot.

![tsne](tsne.png)


<!-- grade_id: cell-2cca477a54e6196e, points: 0, grade: false, solution: true -->
```python
import umap

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f5ba0cf448b8af0a, points: 2, grade: true, solution: false -->
```python
assert(umapD.shape==(731,2))
```