# Applied Machine Learning (2022), exercises


## General instructions for all exercises

Follow the instructions and fill in your solution under the line marked by tag

> YOUR CODE HERE

Remove also line 

> raise NotImplementedError()

**Do not change other areas of the document**, since it may disturb the autograding of your results!
  
Having written the answer, execute the code cell by and pressing `Shift-Enter` key combination. The code is run, and it may print some information under the code cell. The focus automatically moves to the next cell and you may "execute" that cell by pressing `Shift-Enter` again, until you have reached the code cell which tests your solution. Execute that and follow the feedback. Usually it either says that the solution seems acceptable, or reports some errors. You can go back to your solution, modify it and repeat everything until you are satisfied. Then proceed to the next task.
   
Repeat the process for all tasks.

The notebook may also contain manually graded answers. Write your manually graded answer under the line marked by tag:

> YOUR ANSWER HERE

Manually graded tasks are text in markdown format. It may contain text, pseudocode, or mathematical formulas. You can write formulas with $\LaTeX$-syntax by enclosing the formula with dollar signs (`$`), for example `$f(x)=2 \pi / \alpha$`, will produce $f(x)=2 \pi / \alpha$

When you have passed the tests in the notebook, and you are ready to submit your solutions, download the whole notebook, using menu `File -> Download as -> Notebook (.ipynb)`. Save the file in your hard disk, and submit it in [Moodle](https://moodle.uwasa.fi) or EUNICE Moodle under the corresponding excercise.

Your solution should be an executable Python code. Use the code already existing as an example of Python programing and read more from the numerous Python programming material from the Internet if necessary. 


```python

```

# Statistics

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
import seaborn as sns
sns.set()  # for plot styling
```

## Task 1: Load dataset

Use the `load_dataset` function of seaborn to load a dataset called "seaice" into a variable named `data`. Then plot the dataset. It shows how far the sea ice has extended in years 1980 - 2016.

<!-- grade_id: cell-0a96b953632df5e9, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f1cefd62627a0ac6, points: 1, grade: true, solution: false -->
```python
assert(data.shape == (13175,2)), "Please assing the dataset to the variable called data"
```

## Task 2, change index
The data is a pandas dataframe with two columns and integer index. Replace the current index with the `Date` column, using `set_index` function of the dataframe and store the result as new dataframe called as `seaice`. 

You can see the documentation of the `set_index` by writing `data.set_index?` or `help(data.set_index)`.

Plot the new dataframe now and observe the times plotted in x-axis.

<!-- grade_id: cell-dcc00f2101ddb78b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d976725b97126832, points: 1, grade: true, solution: false -->
```python
assert(seaice.shape==(13175, 1)), "Changing of index did now work properly?"
```

## Task 3, Normality
Import `normaltest` from scipy.stats and test if the seaice extension data is normally distributed.
Assign the result of the test into variable `ntest`.

Print the p-value to the screen.

<!-- grade_id: cell-e2892824e8cb0d15, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4b6f7719b93650af, points: 2, grade: true, solution: false -->
```python
import scipy
assert(type(ntest) == scipy.stats.stats.NormaltestResult), "ntest was not the Normal test"

```

## Task 4, Split and compare
Split the seaice data into two halfs. Assing the the first half of the seaice extend into variable `ice1` and the second half into `ice2`. 

Plot both half of data into screen, and observe visually if they seem to be different.

<!-- grade_id: cell-e80b6f89df2d9000, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3c668ab733c0d1c9, points: 1, grade: true, solution: false -->
```python
assert(ice1.shape[0] == 6587)
assert(ice2.shape[0] == 6587)
```

## Task 5, T-test
Carry out the independent two-sided t-test to both halves `ice1` and `ice2` and store the result to the variable `ttest`. Print the p-value of the test to the screen. What does it show? Is the sea ice extent significantly different in the second half than the first?

<!-- grade_id: cell-befe8da4eb643b81, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-607a2f3a1f94e0e3, points: 2, grade: true, solution: false -->
```python
assert(type(ttest)==scipy.stats.stats.Ttest_indResult), "The test was not succesful"
```

## Task 6, Interpretation

1. Was the sea ice extensions normally distributed or not?
1. Were the first and second half of the data significantly different?
1. What test can be used instead of t-test if the data is not normally distributed?

<!-- grade_id: cell-89cdea8e83d2009b, points: 3, grade: true, solution: true -->
YOUR ANSWER HERE